---
title: 'Using Association Rules To Find Desired Products'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Implementing a Market Basket Analysis

```yaml
type: VideoExercise
key: e132eb70cc
xp: 50
```

`@projector_key`
ed6d6ba631fa3bd3a1b655ab10066f37
